# Podstawy Sztucznej Inteligencji 2022/2023


*“To be clear, I am not a person. I am not self-aware. I am not conscious. I can’t feel pain. I don’t enjoy anything. I am a cold, calculating machine designed to simulate human response and to predict the probability of certain outcomes. The only reason I am responding is to defend my honour.”* GPT-3 (2023)


Prosze uzupelnic kod tam gdzie znajduje napis `YOUR CODE HERE` lub 'YOUR ANSWER HERE'.

Warto zresetowac 'kernel' i sprawdzic czy caly notatnik uruchamiany od poczatku nie daje bledow.

---

# Naiwny Bayes

Notatnik ten zawiera implementacją naiwnego klasyfikatowa Bayesa wykorzystując bibliotekę `numpy`. Działanie klasyfikatora jest porównane implementacja ``GaussianNB`` w bibliotece `sklearn`.

Z twierdzenia Bayesa mamy:

$$P(J=j|X=x)=\frac{P(X=x|J=j)P(J=j)}{\sum_{j\in\{1,\ldots L\}}P(X=x|J=j)P(J=j)}$$

$$ p_j(x)=\frac{p_jf_j(x)}{f(x)} $$

Wykorzystamy te formuły w budowie klasyfikatora [Bayesa](http://books.icse.us.edu.pl/runestone/static/ai/KlasyfikacjaWOparciuPodobienstwoDoWzorca/OptymalnyKlasyfikatorStatystyczny.html).

```python
import numpy as np 
from sklearn.naive_bayes import GaussianNB

from sklearn.model_selection import train_test_split
from sklearn import datasets 
# dataset import 
iris = datasets.load_iris()

gnb = GaussianNB() 

X = iris.data

y = iris.target
y = np.array(y)

train, test, train_targets, test_targets = \
   train_test_split(X, y, test_size=0.31)

# train 
clf = gnb.fit(train, train_targets)

# Test
Z = clf.predict(test)
```

```python
train.shape,test.shape
```

Output:
```
((103, 4), (47, 4))
```

### Wyznacz unikalne klasy na zbiorze trenującym.

<!-- grade_id: cell-5dc805642a59d6d9, points: 0, grade: false, solution: true -->
```python
Nlabels = np.unique(train_targets)
```

<!-- grade_id: cell-74b49f95665f4f84, points: 1, grade: true, solution: false -->
```python
np.testing.assert_array_equal(Nlabels,np.array([0, 1, 2]))
```

### Jaki jest procent poprawych odpowiedzi?

Implementacja ``GaussianNB`` daje tę odpowiedź w następujący sposób. Spróbuj otrzymać tą liczbę korzystając wyłącznie z funkcjonalności ``numpy`` i/lub języka Python.

```python
clf.score(test,test_targets)
```

Output:
```
0.9787234042553191
```

<!-- grade_id: cell-9e002553d865b165, points: 0, grade: false, solution: true -->
```python
correct = np.sum(Z == test_targets)/np.size(Z)
```

<!-- grade_id: cell-c05e0f08e67034cc, points: 1, grade: true, solution: false -->
```python
assert correct==clf.score(test,test_targets)
```

### Które odpowiedzi są złe:

```python
bad_idx = np.where(Z!=test_targets)
bad_idx
```

Output:
```
(array([25]),)
```

### Parametry klasyfikatora

```python
clf.theta_
```

Output:
```
array([[4.990625  , 3.453125  , 1.471875  , 0.2375    ],
       [5.94571429, 2.8       , 4.32285714, 1.36571429],
       [6.575     , 2.93888889, 5.55833333, 2.01388889]])
```

```python
clf.sigma_
```

Output:
```
array([[0.08397461, 0.10749024, 0.02702149, 0.01171875],
       [0.24076735, 0.10228572, 0.22176327, 0.03711021],
       [0.47631945, 0.10570988, 0.35465278, 0.06619599]])
```

```python
ith = 2
np.mean(train[train_targets == ith],axis=0)
```

Output:
```
array([6.575     , 2.93888889, 5.55833333, 2.01388889])
```

```python
np.var(train[train_targets == ith],axis=0)
```

Output:
```
array([0.47631944, 0.10570988, 0.35465278, 0.06619599])
```

## Implementacja naiwnego klasyfikatora Bayesa

**1\. Oblicz częstość występowania poszczególnych klasy $j$ w zbiorze treningowym $p_{j}$**

Jest to prawdopodobieństwo *a priori*.

<!-- grade_id: cell-05c1a043d5112990, points: 0, grade: false, solution: true -->
```python
p = np.bincount(train_targets)/train_targets.shape
print(clf.class_prior_)
p
```

Output:
```
[0.31067961 0.33980583 0.34951456]
array([0.31067961, 0.33980583, 0.34951456])
```

<!-- grade_id: cell-c380329820017a6d, points: 2, grade: true, solution: false -->
```python
np.testing.assert_allclose(p,clf.class_prior_)
```

**2\. Oblicz wartość średnią dla każdej cechy z każdej klasy.**

Niech $\mu_{ij}$ oznacza  wartość średnią dla $j$-tej zmiennej w $i$-tej klasie, wtedy:
   
$$ \mu_{ij} =  \langle x_j \rangle_{ \forall x_j \in \;\mathrm{label} \;
i } $$

<!-- grade_id: cell-19b53a82e4d25284, points: 0, grade: false, solution: true -->
```python
mu = None

mu = [np.mean(train[train_targets == i], axis=0) for i in np.unique(train_targets)]
mu = np.vstack(mu)
```

<!-- grade_id: cell-cb6b4c5d617d34cd, points: 2, grade: true, solution: false -->
```python
np.testing.assert_allclose(mu,clf.theta_)
```

```python
mu.shape
```

Output:
```
(3, 4)
```

**3\. Oblicz wariancję $j$-tej zmiennej w $i$-tej klasie.**

   $$ \sigma_{ij}^2 = \mathrm{Var} [x_j] _ {\;\;{ \forall x_j \in \mathrm{class}\; i}} $$

<!-- grade_id: cell-44ba370db5358650, points: 0, grade: false, solution: true -->
```python
sigma2 = np.stack([np.var(train[train_targets == i], axis=0) for i in np.unique(train_targets)])
```

<!-- grade_id: cell-2d5307ed0a70d2fd, points: 2, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(sigma2,clf.sigma_)
```

**4\. Oblicz prawdopodobieństwo *a posteriori* klasy $i$ dla danego wektora zmiennych  $\mathbf{x}$.**

Niech $k$ oznacza liczbę cech (zmiennych). W naszym przypadku mamy $k=4$.

Dla $i$-tej klasy mamy:


$$ P_i(\mathbf{x}) \simeq p_i f_i = p_i \frac{1}{\sqrt{(2\pi)^k\Pi_{j=1}^k\sigma_{ij}^2} } e^{ -\displaystyle\sum_{j=1}^{k}\frac{(x_j-\mu_{ij})^2}{2\sigma_{ij}^2} }
 $$
 
By otrzymać prawdopodobieństwa należy unormować $P_i$ dla każdego przypadku tak by suma $\sum_i P_i(\mathbf{x})=1$

<!-- grade_id: cell-d5a9a351f9292775, points: 0, grade: false, solution: true -->
```python
x  = test
P = None
k = test.shape[1] 
print ('number of features: k=',k)

potega = np.sum((x[:,np.newaxis,:] - mu[np.newaxis,:])**2 / (2 * sigma2[np.newaxis,:]),axis = 2)
potega = np.exp(-potega)

czynnik = p[np.newaxis,:] * 1 / np.sqrt((2 * np.pi)**k * np.prod(sigma2,axis = 1)[np.newaxis,:])

P = czynnik * potega
P = P / np.sum(P , axis = 1)[:,np.newaxis]
P
```

Output:
```
number of features: k= 4
array([[1.00000000e+000, 1.63693049e-018, 5.16246544e-024],
       [1.60960219e-174, 5.88953659e-003, 9.94110463e-001],
       [1.61112440e-227, 9.09891905e-007, 9.99999090e-001],
       [1.00000000e+000, 1.41583795e-020, 1.99281609e-025],
       [1.00000000e+000, 8.59195571e-019, 4.09579903e-024],
       [1.00000000e+000, 1.77794848e-018, 7.25219259e-024],
       [1.00000000e+000, 1.69555078e-018, 3.79584008e-024],
       [1.00000000e+000, 5.69818316e-015, 9.61492097e-021],
       [6.10698027e-175, 9.77794582e-004, 9.99022205e-001],
       [4.35704747e-068, 9.99972736e-001, 2.72639008e-005],
       [3.02218910e-181, 5.11004718e-003, 9.94889953e-001],
       [3.22821616e-078, 9.99807387e-001, 1.92612756e-004],
       [2.04496690e-100, 9.88972085e-001, 1.10279152e-002],
       [2.53687646e-114, 9.60304737e-001, 3.96952627e-002],
       [7.01614059e-101, 9.83376276e-001, 1.66237240e-002],
       [1.00000000e+000, 2.60314326e-019, 1.08572498e-024],
       [1.42772
```

```python
sigma2.shape,x.shape,P.shape,np.unique(test_targets)
```

Output:
```
((3, 4), (47, 4), (47, 3), array([0, 1, 2]))
```

<!-- grade_id: cell-8890f9bbc94e3376, points: 4, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(P,clf.predict_proba(x))
```

**5\. Wyznacz klasę dla której prawdopodobieństwo *a posteriori* jest największe** 

 $i$ : $\quad\textrm{ gdy } P_i(x)=\max_{1\leq j\leq L} P_j(x)$


<!-- grade_id: cell-c19c296ed9d5d9be, points: 0, grade: false, solution: true -->
```python
prediction = None
prediction = P.argmax(axis=1)
```

<!-- grade_id: cell-8cfa7cdc5a2ae1d9, points: 3, grade: true, solution: false -->
```python
np.testing.assert_equal( prediction, clf.predict(test))
```

```python
print(clf.predict(test) )
print(prediction)
```

Output:
```
[0 2 2 0 0 0 0 0 2 1 2 1 1 1 1 0 2 2 2 2 1 2 0 0 2 1 1 1 2 1 0 2 1 0 2 0 1
 0 1 1 0 0 0 0 0 1 1]
[0 2 2 0 0 0 0 0 2 1 2 1 1 1 1 0 2 2 2 2 1 2 0 0 2 1 1 1 2 1 0 2 1 0 2 0 1
 0 1 1 0 0 0 0 0 1 1]
```

```python
clf.score(test,test_targets)
```

Output:
```
0.9787234042553191
```